# Homework \#1

## PHYS 2600, Fall 2022

__Important note on entering text:__  Some of the boxes below ask you to enter some text, with the prompt "_Type your answer here using Markdown._"  We haven't learned any Markdown yet - but that's fine!  Regular text is a subset of Markdown, so you should just type plain text into those cells for your answers.

## 1.1 - Introduction (6 points)

At this level of study in undergraduate physics, students tend to have a pretty wide range of experience in computing, particularly using it for scientific work.  Tell us a little about your background below.  First, do you have any experience with Python or other programming languages?  If so, how much?

<!-- grade_id: intro_background, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Have you taken any computer programming courses before, either here at CU or elsewhere?  If so, what courses did you take?

<!-- grade_id: intro_courses, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

After we get through the basics of Python and some standard numerical techniques, we'll have some room to cover more specific examples.  Is there any one subject, physics application or numerical technique, that you'd be especially interested in seeing?  (I may choose from your answers to this question to plan lectures or assignments late in the semester, so this is your chance to guide the class!)

<!-- grade_id: intro_interests, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## 1.2 - Syntax and Semantics (6 points)

Describe in your own words the difference between _syntax_ and _semantics_ in computer programming.

<!-- grade_id: syntax, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Computers aren't able to flexibly interpret what we tell them to do; they are always perfectly literal.  This is a challenge if you're trying to build a computer to handle a complicated and dynamical process, for example building a self-driving car.

Look up the turn-by-turn driving directions from Boulder to Denver using [Google Maps](http://maps.google.com/).  Now imagine that a self-driving car has been given these directions and nothing more; what are some ways in which the trip could go horribly wrong, compared to what a human driver would do?  Give an example.

<!-- grade_id: self_driving, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## 1.3 - Planetary Weight Conversion (4 points)

Suppose we are standing on an arbitrary planet of mass $M$ and radius $R$.  The force of gravity on us (mass $m$) from the planet is equal to:

\\[
|\vec{F_g}| = \frac{GmM}{R^2}.
\\]

In the code cell below, __implement a formula to calculate the weight (in Newtons) on another planet__ of a person who weighs 650 N on the Earth.  Use it for two planets: Mars (save your answer as a variable named `weight_Mars`) and Mercury (saved as `weight_Mercury`.)  Their mass and radius are:

* Mars: $R = 0.532 R_E$, $M = 0.107 M_E$
* Mercury: $R = 0.383 R_E$, $M = 0.0553 M_E$

The correct answer for the person's weight on Mars is given in the test cell below, to help you check your work.  (Tests are always a good idea in computer programming, to make sure the program does what you expect!)  If your answer for `weight_Mars` is correct, the cell will run without error.

<!-- grade_id: planets, points: 0, grade: false, solution: true -->
```python
weight_Earth = 650  # N

# YOUR CODE HERE
```

<!-- grade_id: planets_tests, points: 4, grade: true, solution: false -->
```python
# Testing cell; don't worry about the details yet, 
# but it should run without error if your answer for Mars in weight_Mars is correct.

import numpy.testing as npt

print("Weight on Earth: %g N" % weight_Earth)
print("Weight on Mars: %d N" % weight_Mars)
print("Weight on Mercury: %d N" % weight_Mercury)

npt.assert_allclose(weight_Mars, 245, rtol=0.01, atol=0)

```

Output:
```
Weight on Earth: 650 N
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_452/4035273970.py in <cell line: 7>()
      5 
      6 print("Weight on Earth: %g N" % weight_Earth)
----> 7 print("Weight on Mars: %d N" % weight_Mars)
      8 print("Weight on Mercury: %d N" % weight_Mercury)
      9 

NameError: name 'weight_Mars' is not defined
```

## 1.4 - Mistakes Were Made (9 points)

All three of the following code examples are implemented incorrectly, containing one or more errors that prevent them from running.  __Fix the code in each example__ to run the correct formula, and give the right answer.

(_Note:_ I emphasize that you have to fix the code and run the correct formula in full - no credit for just entering `print(right answer)`!)

### Part A (3 points)

__Q:__ What is the kinetic energy of a 1500 kg car moving at 12 m/s?

__Formula:__ $KE = \frac{1}{2} m v^2$

__A:__ We should get 108,000 J as the answer.

Copy and paste into the code cell below, and run to see what is wrong.  Then fix the mistakes and run it again.

```python
kinetic-energy = (1/2) * 1500 * v^2

print(kinetic-energy)

```

<!-- grade_id: cell-057f944b5694a43a, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE

```

### Part B (3 points)

__Q:__ An air hockey puck is measured at coordinates x=30 cm, y=30 cm.  How far is it from the origin?

__Formula:__ $d = \sqrt{x^2 + y^2}$

__A:__ We should get 42.4 cm.

Copy and paste into the code cell below, then fix the mistakes.

```python
x = 30
    y = 30

    d = x**2 + y**2**(1//2)

print(d)
```

<!-- grade_id: mistakes_2, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE

```

### Part C (3 points)

__Q:__ A block weighing 800 grams sits on a surface with coefficient of static friction $\mu$ = 0.22.  How much force is required to overcome static friction?  (Take g = 9.8 m/s^2.)

__Formula:__ $F_f = \mu m g$

__A:__ The force required is about 1.72 N.

Copy and paste into the code cell below, then fix the mistakes.

```python
mu = '0.22'
F_N = 800 * 9.8
F_f = mu * F_N

print F_f
```

<!-- grade_id: mistakes_3, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE

```