# Homework #2

## PHYS 2600, Fall 2022

__Important notice:__ All cells in your notebook will be run, start to finish, using a fresh kernel when it is graded!  To make sure the graded notebook looks like what you expect, we recommend selecting "Kernel > Restart & Run All" in the menu above before you finish.



If you worked collaboratively on this assignment, __include the names of your collaborators in the cell below:__

<!-- grade_id: cell-6aad98798978a240, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## 2.1 - Mistakes and formatting (6 points)

Consider the following program, which tries to verify that these two equations are correct:

$$
(a+b)^2 = a^2 + 2ab + b^2 \\
(a-b)^2 = a^2 - 2ab + b^2
$$

for the values $a=3.3$ and $b=5.3$.

```python
a=3,3   b=5,3
a2 = a**2
b2 = b**2

eq1_sum = a2 + 2ab + b2
eq2_sum = a2 - 2ab - b2

eq1_pow = (a + b)**2
eq2_pow = (a - b)**2

print(f'First equation: {eq1_sum:g} = {eq1_pow:g}')
print(f'Second equation: {eq2_sum:h} = {eq2_pow:h}')
```

There are tons of mistakes in this program!  Copy the code into the cell below, then __fix the mistakes__ and run to verify the two equations.

<!-- grade_id: cell-73dc6f433ff4bab7, points: 6, grade: true, solution: true -->
```python

a = 3.3   
b = 5.3
a2 = a**2
b2 = b**2

eq1_sum = a2 + 2 * a * b + b2
eq2_sum = a2 - (2 * a * b) + b2

eq1_pow = ((a + b) ** 2)
eq2_pow = ((a - b) ** 2)

print(f'First equation: {eq1_sum:g} = {eq1_pow:g}')
print(f'Second equation: {eq2_sum:g} = {eq2_pow:g}')

```

Output:
```
First equation: 73.96 = 73.96
Second equation: 4 = 4
```

## 2.2 - Dealing with multiple arguments (6 points)

This problem concerns some typical 'gotcha' behaviors of function arguments that can lead to potential confusion.  Here is a brief reminder of some behaviors of functions that we discussed:

* Positional arguments given to a function are assigned in order to the argument names as in the `def...` statement.
* Keyword arguments can be in _any_ order, and are assigned to matching argument names as in the `def...` statement.
* A function call has its own __local scope__; any variable names inside are temporary and distinct from those in the original (global) scope.

For each of the code examples below, __explain the question in the comment__.  If you get stuck, refer back to lecture and tutorial 6, and don't be afraid to use the [Python Tutor](http://www.pythontutor.com/visualize.html#mode=edit)!

```python
# Example A
x = 2
def H(x):
    print(x)
H(x=3)
print(x)  ## Why doesn't this print 3? 
```

Output:
```
3
2
```

<!-- grade_id: cell-c5263d42e84d426f, points: 2, grade: true, solution: true -->
Because this x uses the global variable x nd not the local H variable x

```python
# Example B
x = 2; y = 3
def F(y,x):
    print(y,x)

F(x, y)  ## Why does this print 2 3, and not 3 2?
```

Output:
```
2 3
```

<!-- grade_id: cell-483c01c7994d4b8e, points: 2, grade: true, solution: true -->
Because you put the global variable x in the y position of the function, assigning it the value 2, and put the global variable y in the x position, assigning it the value 3.

```python
# Example C
x = 2; z = 3
def G(x, y=1, z=2):
    print(x,y,z)

G(x, z)  ## Why does this print 2 3 2, and not 2 1 3?
```

Output:
```
2 3 2
```

<!-- grade_id: cell-ccf321d083974b26, points: 2, grade: true, solution: true -->
G is using the global variable z which has the value 3. Because it is the second input, the computer sees it in the position of the y input in G (the second input), and G will default the third input (z) to what is defined: 2. 

## 2.3 - Strings in the lab (5 points)


You're working in a physics lab, recording some data from a tabletop experiment.  The experiment produces three numbers: a time-of-flight `t` in seconds, a polarization angle between $(-\pi, \pi)$ in radians, and a light frequency in Hz.

In order to do further analysis on the raw data, your research advisor has asked you to write them out to a data file.  However, they want the data formatted in a very specific way:

- The time should be in milliseconds, not seconds, and should be printed out to the nearest _micro_second (i.e. with three digits after the decimal point);
- The angle should be shifted to be between $(0, 2\pi)$ instead of $(-\pi, \pi)$, and then should be printed to three digits after the decimal;
- The frequency should be printed in scientific notation, with four significant figures.

Given the following measurements:

- Time of flight = 0.00418834 s
- Polarization angle = -0.1934 rad
- Light frequency = 49120930018493 Hz

__Make a string called `format_data`__ which contains the three numbers with formatting matching your advisor's request.  The three numbers should be printed in the order "(time) (angle) (frequency)", with a single space separating them.

(So far we don't know how to actually write this out to a file - we'll cover that later - but if you were really doing this task in a lab, this formatting exercise would still be the first step!)

<!-- grade_id: cell-bf158df8dd0c470d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from math import pi

time = 0.00418834 # s
angle = -0.1934 # rad
frequency = 49120930018493 # Hz
def format_data_function(time, angle, frequency):
   time_milli = time * 10**-3
   angle_pi = angle + pi

   format_data = '{time_milli}, {angle_pi}, {frequency:.4g}'.format(time_milli=time_milli, angle_pi=angle_pi, frequency=frequency)
   return format_data

format_data = format_data_function(time, angle, frequency)
   

```

<!-- grade_id: cell-631934b45b48d8e5, points: 5, grade: true, solution: false -->
```python
# Testing cell - do not delete me!

print(format_data)
assert type(format_data) == str

# NOTE: the visible tests only make sure you created the string format_data.
# Just because this cell runs without error, it doesn't mean your answer is right
# - check it yourself against the requirements!!

```

Output:
```
4.18834e-06, 2.948192653589793, 4.912e+13
```

## 2.4 \- Simple projectile motion \(5 points\)



From introductory mechanics, we know that the formula for the height of a ball which is thrown into the air with initial vertical speed $v_0$ is (ignoring air resistance):

$$
y(t) = y_0 + v_0 t - \frac{1}{2} gt^2
$$

with $g = 9.8$ m/s${}^2$.  Going one step further, we expect that at some later time $t_f$, the ball will hit the ground.  The condition for this is, taking the ground to be at zero height, $y(t_f) = 0$.  We can use this to solve for $t_f$, finding

$$
t_f = \frac{v_0 + \sqrt{v_0^2 + 2gy_0}}{g}
$$

In the cell below, __implement a function called `time_of_flight()`.__  It should have the following features (this might be called a _specification_, which is a sort of contract as to what a function yet to be written should look like):

- It should have one __required__ argument (positional), called `v_0`.
- It should have two __optional__ arguments, in any order: `y_0` (default value: 0.0) and `g` (default value: 9.8.)
- It should return `t_f` according to the formula above.

To compute the square root, you can either import the `sqrt` function from the `math` module, or you can raise to the $0.5$ power.  (The `sqrt` function is much more efficient, but for a simple task like this, you won't notice any difference in speed!)

<!-- grade_id: cell-9e55289b55a580d6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from cmath import sqrt

y_0 = 0.0
g = 9.8

def time_of_flight(v_0, y_0, g):
    t_f = (v_0 + sqrt(v_0**2 + 2 * g * y_0)) / g

    return t_f
```

<!-- grade_id: cell-9e65a120ef01a246, points: 5, grade: true, solution: false -->
```python
## Testing cell

import numpy.testing as npt

npt.assert_allclose(time_of_flight(3.2), 0.653, atol=1e-3)
npt.assert_allclose(time_of_flight(3.2, g=1.63), 3.926, atol=1e-3)  # Moon gravity instead of Earth
npt.assert_allclose(time_of_flight(0.0), 0.0, atol=1e-3)  # y_0 = 0 and v_0 = 0 gives t_f = 0

```

Output:
```
Canceled future for execute_request message before replies were done
The Kernel crashed while executing code in the the current cell or a previous cell. Please review the code in the cell(s) to identify a possible cause of the failure. Click <a href='https://aka.ms/vscodeJupyterKernelCrash'>here</a> for more info. View Jupyter <a href='command:jupyter.viewOutput'>log</a> for further details.
```

## 2.5 - Intro to Markdown (8 points)

Now that we've learned how to print things out nicely in Python, let's dive into the other way of writing nice text in Python, the Markdown markup language.  Markdown is ubiquitous on the internet; in particular, Wikipedia and other wikis use Markdown natively for notation.  WhatsApp uses a subset of Markdown for formatting, so you may have encountered it there.  And of course, almost all the class materials are written in Markdown!

We'll start with the basics of Markdown notation, and cover math notation using MathJax another time.  A more complete Markdown cheat sheet is [available here](https://github.com/adam-p/markdown-here/wiki/Markdown-Cheatsheet).

Start by double-clicking on the Markdown cell immediately below this one to see the Markdown source.  The actual problem appears in the cell after the big example.

# Headings
## Come
### In
#### Five
##### Sizes
###### (Sometimes six, but not in Jupyter)

## Text emphasis

We can add _emphasis_ to our text in a few different ways: _italic_, __bold__, or ~~strikethrough~~.  There are alternative ways to write *italic* and **bold** - we can mix them together to get _**bold** italics_.

## Code

Often we want to write Python code snippets in our Markdown blocks, which uses a special notation to get the gray highlighting and different font (called 'monotype' since the characters are all the same width): `print("Hello, world!")`

We can put multi-line code snippets in our Markdown, too:
```
x=1
y=2
print(x**y)
```
If we add the language name to the top, we can even get nice _syntax highlighting_, just like Jupyter does for us in the code cells!
```python
x=1
y=2
print(x**y)
```

## Lists

Markdown is great at making both ordered and unordered lists in a natural way!  Write the ordered list with numbers:

1. Total energy is always conserved.
9. Entropy of an isolated system always increases.
5. Entropy of an isolated system approaches a universal constant as T goes to absolute zero.

Slightly counter-intuitive: Markdown doesn't care which numbers you use, or how you order them - it gives you a simple numbered list!  Unordered lists can be written with asterisks:

* Bread
* Eggs
* Milk
+ Yogurt
- Bananas


Dashes and plus signs are equally valid for unordered lists; you can even mix and match (as seen above.)

## External links

There's a complicated set of rules for adding all kinds of links to a Markdown document, but the simplest useful one is just linking to something external.  The syntax has two entries: the first is the _label_, and the second is the _link_.

[This is a link to Google.](http://www.google.com)

The link itself should be a valid URL - exactly what appears in your web browser bar when you load a page.  You can just copy directly out of the browser bar to grab a page link that you're looking at.

__Now the homework problem!__  Your mission is to give a brief explanation in Markdown, in your own words, of the principle of conservation of energy.  (If you need to jog your memory, this page might be helpful:)

[https://cnx.org/contents/Ax2o07Ul@13.5:xa2wEvur@9/Conservation-of-Energy](https://cnx.org/contents/Ax2o07Ul@13.5:xa2wEvur@9/Conservation-of-Energy)

To fully flex your Markdown muscle, __you must include the following in your explanation:__

- One first-level header reading "Conservation of Energy";
- Three second-level headers, reading "The Law of Conservation of Energy", "Types of Energy", and "Examples of Energy Conversion".
- At least one example of each type of emphasis: bold, italic, and bold italic.
- A list, ordered or unordered, containing three different examples of types of energy (your first can be "kinetic energy".)
- An external link to the URL given above, with some text description.
- A multi-line code example, showing a Python snippet which adds together the kinetic energy and gravitational potential energy for a 1 kg mass moving at 5 m/s, and at a height of 12 meters.




<!-- grade_id: cell-4ebb5e74ff0b0e94, points: 8, grade: true, solution: true -->
# Conservation of Energy

## The Law of Conservation of Energy
The total energy must remain the same of a closed system must remain the same.

## Types of Energy
- Kinetic
_Ex:_ **a moving truck**
- Potential
_Ex:_ **gravitational**, _**voltage**_

[Read More] (https://openstax.org/books/college-physics/pages/7-6-conservation-of-energy)

## Code
```python 

m = 1 #kg
v = 5 #m/s
h = 12 #m

def total_energy:
    g = 9.8 #m/s**2
    KE = 0.5 * m * v**2
    GPE = m * g * h
    E = KE + GPE
    
    return E
```