# Tutorial 4: Functions and Scope

## PHYS 2600, Fall 2022

## T4.1 - A deep dive with variables and scope

_(Special note: part A of this first problem is a __worked example__, which we'll go through together in class.  You are encouraged to fill this in as you follow along, but you won't be graded on whether you've completed it or not.)_

### Part A

Let's dig into a more complicated example of a function definition, and see what's happening with variables and their scope, line by line.  We'll make the function itself really simple: it will be called `triple(x)`, and will return 3 times whatever `x` is.  

Consider the following code, but __do not run it yet:__

```python
x=2
y=3

def triple(x):
    print("Inside triple, x is:", x, "and y is:", y)
    z=3*x
    return z

print("Triple sum:", triple(3) + triple(4))
print("At the end, x is:", x)
```

Output:
```
Inside triple, x is: 3 and y is: 3
Inside triple, x is: 4 and y is: 3
Triple sum: 21
At the end, x is: 2
```

Now, let's make sure we understand how variables and scope work by __predicting what will happen in the code, step by step.__  In class, I'll do this (with your help!) at the blackboard.  If you're working this exercise at home, make a blank Markdown cell or use a sheet of paper to predict what the code will do.  

For context, here's an example sketch of what the global and local namespaces look like after we call `triple(4)`:

<img src="https://physicscourses.colorado.edu/phys2600/phys2600_sp19/img/scopes.png" width=500px />

When you're finished with your prediction, run the cell and see if the output is what you expected.  Then __check your prediction__ by running the code above through the [Python Tutor](http://www.pythontutor.com/visualize.html#mode=edit).

__For the rest of this problem, continue on your own.__

### Part B

As the diagram hints, we can access global scope from inside a function as well, as long as we don't override the variable name.  __Predict the output of the following example__, and then run it:

```python
n=2

def global_pow(x):
    return x**n

print(global_pow(3))

n=3
print(global_pow(3))
```

__Why does changing `n` work like this__, if we do it _after_ we've defined the `global_pow` function already?!  (The answer has to do with _when_ the statement `return x**n` is actually evaluated; again, try the Python Tutor if you're stuck.)

<!-- grade_id: cell-a5994a4458c6db4f, points: 0, grade: true, solution: true -->
global_pow uses the last define n value.

You should keep in mind that using global variables in this way is almost always a bad idea!  There are two major reasons:

1. The code is much less clear; to know what `global_pow()` will actually do, you have to go backwards to find where `n` was defined last.
2. If $n$ is used somewhere else in your code, the behavior of `global_pow()` can suddenly and unexpectedly change!

### Part C

Since using global variables has those drawbacks, let's implement a better `pow` function using keyword arguments - which will reveal another surprising behavior!  Again, __predict the output of the following code__, then run it:

```python
n=2

def keyword_pow(x, n=n):
    return x**n

print(keyword_pow(3))
print(keyword_pow(3,n=5))

n=3
print(keyword_pow(3))
print(keyword_pow(3, n=n))
```

Output:
```
9
243
9
27
```

Once again, if you get stuck the Python Tutor will help.  Answer the following questions:

- What does the bizarre-looking statement `n=n` actually do?
- Why doesn't the behavior of `keyword_pow(3)` change after we change the global value of `n` to 3?

<!-- grade_id: cell-2e5e642b0d1b240b, points: 0, grade: true, solution: true -->
It uses the global defined value of n into the keyword_pow function.
Because keyword_pow works off of its locally defined n.

## T4.2 - Making your own functions

Let's start with some basic practice making our own functions.  Each of the code examples below contains a function call for a function that hasn't been implemented yet.  In each part below, __write the function in the first cell__ so that the code in the second cell works correctly.

### Part A

We want a function called `hello` that takes no arguments, and returns the string `Hello, world!` when it is called.

<!-- grade_id: cell-764aff9929c55d48, points: 2, grade: true, solution: true -->
```python
# Write your function here

def hello():
    return 'Hello, world!'
```

```python
## Read-only code: should run after you implement the function above.

print(hello())

# Note: this should print only "Hello, world!" - if you see "None" printed out as well, your function is incorrect!
# Remember that you should _return_ the string from your function.
```

Output:
```
Hello, world!
```

### Part B

We want a function called `dist_origin` that takes two numbers `x` and `y`, and computes the distance from the origin $(0,0)$ to the point $(x,y)$.  (You'll need to import a square-root function from somewhere!)

<!-- grade_id: cell-71a3f74b37afb80a, points: 0, grade: false, solution: true -->
```python
# Write your function here

from cmath import sqrt


def dist_origin(x,y):
    return sqrt(x**2+y**2)
```

```python
## Read-only code: should run after you implement the function above.

print(dist_origin(3,4))  # should print 5.0
print(dist_origin(1,-1)) # should print about 1.41
```

Output:
```
(5+0j)
(1.4142135623730951+0j)
```

### Part C

A generalization of the distance function $d = \sqrt{x^2 + y^2}$ in mathematics is the __p-norm__, which is defined as follows for a two-dimensional vector:

$$
d_p \equiv (|x|^p + |y|^p)^{1/p}.
$$

For $p = 2$ this is just the usual distance to the origin; if we take $p=1$, it is the sum of the lengths along the $x$ and $y$ directions (the "taxicab distance".)

Write a function called `p_norm_dist()` that takes two __positional arguments__, `x` and `y`, and a single __keyword argument__ `p` (with default value 2).

(_Hint: don't forget the absolute values!  You can use the built-in `abs()` function._)

<!-- grade_id: cell-4c597254fa56b258, points: 0, grade: false, solution: true -->
```python
# Write your function here

p=2
def p_norm_dist(x,y,p=p):
    return (abs(x)**p + abs(y)**p)**(1/p)
```

```python
## Read-only code: should run after you implement the function above.

print(p_norm_dist(3,4))          # Should print 5.0 again
print(p_norm_dist(3,4,p=1))      # Should print 7.0
print(p_norm_dist(x=1,y=-1))     # Should print 1.41... again
print(p_norm_dist(x=1,y=-1,p=1)) # Should print 2.0

print(p_norm_dist(7,3,p=100))    # Should print 7.0
```

Output:
```
5.0
7.0
1.4142135623730951
2.0
7.0
```

### Part D

Whenever we're using a function, and _especially_ if it's imported from a module, it is important to know the __call signature__ of that function.  In short:

- What are the arguments called?
- What order are they in?
- Which arguments have default values (and what are they?)

In Python, the call signature is simply the header line.  But if we're importing from a module, we don't want to have to go read that module's source code!

Fortunately, Jupyter provides a better way, through something called a __magic command__.  Magic commands (or "Jupyter magic") are special instructions that are not part of Python, but are understood by Jupyter.

For dealing with functions, the `?` magic command is essential.  Try it on the function you just wrote: in the cell below, type `p_norm_dist?` and look for the pop-up at the bottom of the screen!

Then give it a try on another module: enter `math.log?`.  (If you get an error, make sure `math` has been imported in your current kernel.)

<!-- grade_id: cell-d81b0fec5df3bb39, points: 0, grade: false, solution: true -->
```python
import math
math.log?

```

Output:
```
[0;31mDocstring:[0m
log(x, [base=math.e])
Return the logarithm of x to the given base.

If the base not specified, returns the natural logarithm (base e) of x.
[0;31mType:[0m      builtin_function_or_method
```

## T4.3 - Mixing scope with the `global` keyword

There's one question I never addressed about scope in the materials so far: can you _write_ to global variables from inside of a local scope?  The answer is sort of, but it leads to one of the most complicated parts of Python's scoping rules.  For your information, we'll go through that case now.

This section comes with a caveat: _you should almost never use the `global` keyword!_  Just as is the case with the global variable example we saw in lecture, modification of global variables breaks the principle of __encapsulation__, that a function should just be concerned with things locally available to it.  Encapsulation makes it much easier to diagnose errors; if your code depends on global scope, you have to look _everywhere_ for the problem!

### Part A


Although the variables in global namespace are visible within the scope of a function, Python actually distinguishes between _reading_ and _writing_ such variables.  Consider the following example code:

```python
n=2

def adder():
    global n 
    n += 1

adder()
print(n)
adder()
print(n)
```

Output:
```
3
4
```

First, read the code and make sure you understand what it's supposed to do, and what the expected outcome of the final `print(n)` is.  Then, run the code and notice that you get an error message, specifically an `UnboundLocalError`.  This is because Python's default behavior is to _not_ allow writing to global variables inside of local scope.

However, we can overrule this behavior by using the `global` keyword.  The statement `global x` tells Python that we want to be able to make full use of the global-namespace variable `x` in our local scope.  __Use the `global` keyword to fix the code above__, so that `adder()` has the expected outcome.

### Part B

One more example of using global scope.  The _Fibonacci sequence_ is defined by the formula
\\[
F_n = F_{n-1} + F_{n-2}
\\]
along with the definitions $F_0 = F_1 = 1$.  Let's make use of global scope to compute numbers from this sequence.  In the cell below, implement a function `iterate_fib()` which uses the values of the global variables `Fn_minus_one` and `Fn_minus_two` to compute the next value in the sequence.  

`iterate_fib()` should __return__ the value $F_n$, and then change `Fn_minus_one` and `Fn_minus_two` to prepare for the next time the function is called.  Remember that the `global` keyword is needed if you're going to assign to variables outside the local scope in a function.

<!-- grade_id: cell-0435e4dadc8a037c, points: 0, grade: false, solution: true -->
```python
Fn_minus_one = 1
Fn_minus_two = 1

def iterate_fib():
    global Fn_minus_one
    global Fn_minus_two

    Fn = Fn_minus_one + Fn_minus_two
    Fn_minus_two = Fn_minus_one
    Fn_minus_one = Fn
   
    
    
    return Fn

    # YOUR CODE HERE
```

<!-- grade_id: cell-38f50cde44d543b8, points: 0, grade: true, solution: false -->
```python
## Testing cell; should run without error if your implementation above works!

Fn_minus_one = 1
Fn_minus_two = 1

assert (iterate_fib() == 2)
assert (iterate_fib() == 3)
assert (iterate_fib() == 5)
assert (iterate_fib() == 8)
assert (iterate_fib() == 13)
```

```python

```